### Instructions:

- You can attempt any number of questions and in any order.  
  See the assignment page for a description of the hurdle requirement for this assessment.
- You may submit your practical for autograding as many times as you like to check on progress, however you will save time by checking and testing your own code before submitting.
- Develop and check your answers in the spaces provided.
- **Replace** the code `raise NotImplementedError()` with your solution to the question.
- Do **NOT** remove any variables other provided markings already provided in the answer spaces.
- Do **NOT** make any changes to this notebook outside of the spaces indicated.  
  (If you do this, the submission system might not accept your work)

### Submitting:

1. Before you turn this problem in, make sure everything runs as expected by resetting this notebook.    
   (You can do this from the menubar above by selecting `Kernel`&#8594;`Restart Kernel and Run All Cells...`)
1. Don't forget to save your notebook after this step.
1. Submit your .ipynb file to Gradescope via file upload or GitHub repository.
1. You can submit as many times as needed.
1. You **must** give your submitted file the **identical** filename to that which you downloaded without changing **any** aspects - spaces, underscores, capitalisation etc. If your operating system has changed the filename because you downloaded the file twice or more you **must** also fix this.  



---

### [Shubharthak Sangharasha](https://shubharthaksangharsha.github.io)
### Student ID: a1944839
### CompScience Foundation-A Assignment-1

#### Question 1&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(2 Point)
Write a Python program to calculate the length of the sentence `I love Python Programming` using string methods. Create a variable `output1` and assign the length of the sentence to it.

For example, if the text was `I love Lucy` you would calculate the length of the string in your program and assign that value to `output1`.


<!-- grade_id: question_1, points: 0, grade: false, solution: true -->
```python
txt = "I love Python Programming"

# Write your solution here
output1 = 0 #alternative solution using for loop
for _ in txt:
    output1 += 1
print(f"Length of string(using for loop) is: {output1}")    
output1 = len(txt) #solution1 - using string method 
print(f"Length of string((using string method) is: {output1}")

```

Output:
```
Length of string(using for loop) is: 25
Length of string((using string method) is: 25
```

<!-- grade_id: 1-1_String_Length, points: 2, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
assert txt == "I love Python Programming", 'txt should be "I love Python Programming"'
```

#### Question 2&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(2 Point)
Write a Python program to print a random integer value in the range of 1 to 10. Create a variable `output2` and assign the solution to it.

<!-- grade_id: question_2, points: 0, grade: false, solution: true -->
```python
import random

# Write your solution here
output2 = random.randint(1, 10) #using random.randint (it will give a random integer [a, b] where a and b are included)
# YOUR CODE HERE
print(f"Random integer between 1-10 is {output2}")
```

Output:
```
Random integer between 1-10 is 8
```

<!-- grade_id: 2-1_Random_Integer, points: 2, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 3&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(2 Point)
Write a Python program to print the data type of -35.59. Create a variable `output3` and assign the solution to it.

<!-- grade_id: question_3, points: 0, grade: false, solution: true -->
```python
y = -35.59

# Write your solution here
output3 = type(y) #type function will provide the datatype of the variable
# YOUR CODE HERE
print(f"Datatype of {y} is {output3}") #prints float
```

Output:
```
Datatype of -35.59 is <class 'float'>
```

<!-- grade_id: 3-1_Data_Type, points: 2, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 4&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(2 Point)
Write a Python program to get the answer of 4<sup>20</sup>. Create a variable `output4` and assign the solution to it.

<!-- grade_id: question_4, points: 0, grade: false, solution: true -->
```python
# Write your solution here
output4 = 4 ** 20 # ** is the power/exponential operator which gives a**b meaning a to power b
# YOUR CODE HERE
print(f"4 raise to power 20 is {output4}")
```

Output:
```
4 raise to power 20 is 1099511627776
```

<!-- grade_id: 4-1_Exponents, points: 2, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 5&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(2 Point)
Define one string as the word "Problem" and define another string as the word "Solving". Combine these two strings to make the statement "Problem Solving". Create a variable `output5` and assign the solution to it.

<!-- grade_id: question_5, points: 0, grade: false, solution: true -->
```python
# Write your solution here
string1, string2 = "Problem", "Solving"
output5 = string1 + " " + string2
# YOUR CODE HERE
print("Combination of string1 and string2 using string concatenation is : {output5}")
```

Output:
```
Combination of string1 and string2 using string concatenation is : {output5}
```

<!-- grade_id: 5-1_String_Concatenation, points: 2, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 6&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(3 Points)
Set the values 7 and 3 to the variables `x` and `y` (You can do this in one line in Python, if needed - see Hints).  
What is the value & data type of `x/y`?. Create a variable `output6_i` and assign the **data type** to it.  
What is the value & data type of `x//y`?. Create a variable `output6_ii` and assign the **data type** to it.

***
<details>
  <summary><span style="color:blue">Click to see Hints</span></summary>
    
* Python allows you to assign values to multiple variables in one line. e.g., `name1, name2 = "John", "David"`
</details>

<!-- grade_id: question_6, points: 0, grade: false, solution: true -->
```python
# Write your solution here
x, y = 7, 3 
output6_i = type(x / y)
output6_ii = type(x // y)
# YOUR CODE HERE
print(f"x/y = {x/y}, datatype is {output6_i}")
print(f"x//y = {x//y}, datatype is {output6_ii}")
```

Output:
```
x/y = 2.3333333333333335, datatype is <class 'float'>
x//y = 2, datatype is <class 'int'>
```

<!-- grade_id: 6-1_Integer_Division, points: 3, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
assert x == 7, 'x should be 7'
assert y == 3, 'y should be 3'
```

#### Question 7&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(3 Points) 
Pick a number `n` to make the following statement True: (5<=n). Create a variable `output7` and assign the solution to it to show your number works.

<!-- grade_id: question_7, points: 0, grade: false, solution: true -->
```python
# Write your solution here
n = 5
output7 = 5 <= n
# YOUR CODE HERE
print(f"5 <=n: {output7}")
```

Output:
```
5 <=n: True
```

<!-- grade_id: 7-1_Boolean_Expression, points: 3, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 8&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(3 Points) 
Write a Python program to convert the string variable `user_height` with value `"175"` into an integer. Create a variable `output8` and assign the solution to it.

```python
user_height = "175"
```

<!-- grade_id: question_8, points: 0, grade: false, solution: true -->
```python
# Write your solution here
output8 = int(user_height) #converting string to int using typecasting 
# YOUR CODE HERE
print(f"Converting {user_height} {type(user_height)} to int: {output8}, {type(output8)}")
```

Output:
```
Converting 175 <class 'str'> to int: 175, <class 'int'>
```

<!-- grade_id: 8-1_Type_Conversion, points: 3, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 9&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(3 Points)
Create a variable `sentence` and assign the string literal `John said "Hello"!` to it. Create a variable `output9` and assign the number of occurrences of double quotes (") in `sentence` using string methods.

<!-- grade_id: question_9_setup, points: 0, grade: false, solution: true -->
```python
# Create your variable here
sentence = 'John said "Hello"!'
# YOUR CODE HERE
print('Sentence:', sentence)
```

Output:
```
Sentence: John said "Hello"!
```

Now create a variable `output9` and using string methods find the number of occurrences of double quotes (") in `sentence`. Assign this number of occurrences to `output9`.

<!-- grade_id: question_9, points: 0, grade: false, solution: true -->
```python
# Write your solution here
output9 = sentence.count('"')
# YOUR CODE HERE
print(f'Occurences of Double Quotes ("") is {output9}')
```

Output:
```
Occurences of Double Quotes ("") is 2
```

<!-- grade_id: 9-1_String_counts, points: 3, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 10&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(4 Points)
Write a Python program to convert a given temperature in Fahrenheit (F) to Celsius (C).  

First create a variable `fahrenheit_temperature` and assign it a Fahrenheit temperature value.

<!-- grade_id: question_10_setup, points: 0, grade: false, solution: true -->
```python
# Create your variable here
fahrenheit_temperature = 45.8

# YOUR CODE HERE
print("fahrenheit temperature - {}".format(fahrenheit_temperature))
```

Output:
```
fahrenheit temperature - 45.8
```

Use the following formula to convert the Fahrenheit value to Celcius.

$C=(F-32)\times \frac{5}{9}$

Create a variable `output10` and assign the final value to it **as a whole number**.

<!-- grade_id: question_10, points: 0, grade: false, solution: true -->
```python
# Write your solution here
celsius_temperature = (fahrenheit_temperature - 32) * 5 / 9
output10 = round(celsius_temperature)
# YOUR CODE HERE
print(f"{fahrenheit_temperature}°F = {output10}°C")
```

Output:
```
45.8°F = 8°C
```

<!-- grade_id: 10-1_Temperature_conversion, points: 4, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 11&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(5 Points)
Write a Python program that uses python's built-in methods to format a given string so that the first letter of each word is made uppercase (and the other letters are lowercase).

For example, "thE universiTy OF adelaiDE" would change to "The University Of Adelaide". 

Create a variable `output11` and assign the solution to it.

***

<details>
  <summary><span style="color:blue">Click to see Hints</span></summary>
    <ul><li>
    Review the <a href="https://docs.python.org/3/library/stdtypes.html#string-methods" >python built-in types documentation for String methods</a> for methods that do some or all of these steps.</li>
        <li>This type of formatting is also called <em>titlecase</em>.</li>
    </ul>
</details>

```python
input11 = "thE universiTy OF adelaiDE"
```

<!-- grade_id: question_11, points: 0, grade: false, solution: true -->
```python
# Write your solution here
output11 = input11.title() #title() inbuild-python function to make first alphabet as upperacse and rest as lowercase
# YOUR CODE HERE
print(f"Answer: {output11}")
```

Output:
```
Answer: The University Of Adelaide
```

<!-- grade_id: 11-1_String_Formatting, points: 5, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 12&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(4 Points)
The sentence `Python 3.10.1, documentation released on 6 December 2021` contains a mixture of alphabetical characters (a-z), numbers (0-9), symbols (., ).  

Write a Python program to check the input for existence of alphabetical characters, numbers and symbols as described below using string methods.
1. Check if all the characters in `input12` are letters. Create a variable `output12_i` and assign the solution to it.
2. Check if all the characters in `input12` are numbers. Create a variable `output12_ii` and assign the solution to it.
3. Check if `input12` only contains alphanumeric characters - letters and numbers (example of characters that are not alphanumeric are (space)! # % & ? etc.). Create a variable `output12_iii` and assign the solution to it.

Try with different input values, such as `"Hello"`, `"123456"`, `"H3ll0"`, and `"H3ll0 W0rld"`

***

<details>
  <summary><span style="color:blue">Click to see Hints</span></summary>
    <ul><li>
    Review the <a href="https://docs.python.org/3/library/stdtypes.html#string-methods" >python built-in types documentation for String methods</a> for methods that do some or all of these steps.</li>
        <li>Built-in methods that check for some condition (such as only alphabetic characters) usually start with <em>is</em>, for example <code>isupper()</code> which checks that a string is all uppercase.</li>
    </ul>
</details>

```python
input12 = "Python 3.10.1, documentation released on 6 December 2021"
```

<!-- grade_id: question_12, points: 0, grade: false, solution: true -->
```python
# Write your solution here
#Checking if it contain letters 
output12_i = input12.isalpha()
#Checking if it contains numbers 
output12_ii = input12.isdigit()
#Checking if it contains alphanumeric characters
output12_iii = input12.isalnum()

#Trying with ifferent values 
input12_eg1 = "Hello"
input12_eg2 = "123456"
input12_eg3 = "H3ll0"
input12_eg4 = "H3ll0 W0rld"

print(f"Checking for {input12_eg1} - {input12_eg1.isalpha()}")
print(f"Checking for {input12_eg2} - {input12_eg2.isdigit()}")
print(f"Checking for {input12_eg3} - {input12_eg3.isalnum()}")
print(f"Checking for {input12_eg4} - {input12_eg3.isalnum()}")

# YOUR CODE HERE
print(f"Output I (All letters) for {input12}: {output12_i}")
print(f"Output II (All numbers) for {input12}: {output12_ii}")
print(f"Output III (Alphanumeric only) for {input12}: {output12_iii}")
```

Output:
```
Checking for Hello - True
Checking for 123456 - True
Checking for H3ll0 - True
Checking for H3ll0 W0rld - True
Output I (All letters) for Python 3.10.1, documentation released on 6 December 2021: False
Output II (All numbers) for Python 3.10.1, documentation released on 6 December 2021: False
Output III (Alphanumeric only) for Python 3.10.1, documentation released on 6 December 2021: False
```

<!-- grade_id: 12-1_String_checks, points: 4, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 13&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(4 Points)
Euler's number, `e`, can be called in Python using `from math import e`,  
1. Round `e` to the nearest integer. Create a variable `output13_i` and assign the solution to it.
2. Round `e` to the 2 decimal places. Create a variable `output13_ii` and assign the solution to it.
3. Round `e` to the 1 decimal place. Create a variable `output13_iii` and assign the solution to it.

<!-- grade_id: question_13, points: 0, grade: false, solution: true -->
```python
from math import e

# Write your solution here
output13_i = round(e) #rounding to nearest integer
output13_ii = round(e, 2) #rounding to 2 decimal places 
output13_iii = round(e, 1) #rounding to 1 decimal places

# YOUR CODE HERE
print(f"Rounding to nearest integer: {output13_i}")
print(f"Rounding to 2 decimal places: {output13_ii}")
print(f"Rounding to 1 decimal places: {output13_iii}")
```

Output:
```
Rounding to nearest integer: 3
Rounding to 2 decimal places: 2.72
Rounding to 1 decimal places: 2.7
```

<!-- grade_id: 13-1_E, points: 4, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 14&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(5 Points)
Write a program that removes all empty spaces, commas and full stops from a given input string using string methods. Create a variable `output14` and assign the solution to it.

For example, `  .i., .nc.,., .omp  r.,ehe. ns,, .,. ibi . ,l,iti,., es ` should output as _incomprehensibilities_

***

<details>
   <summary><span style="color:blue">Click to see Hints</span></summary>
<ul><li>
    Review the <a href="https://docs.python.org/3/library/stdtypes.html#string-methods" >python built-in types documentation for String methods</a> for methods that do some or all of these steps.</li>
        <li>You can remove specific characters from a string by finding and replacing them with an empty string <code>""</code>.</li>
    </ul>
</details>

```python
input14 = "  .i., .nc.,., .omp  r.,ehe. ns,, .,. ibi . ,l,iti,., es "
```

<!-- grade_id: question_14, points: 0, grade: false, solution: true -->
```python
# Write your solution here
#we can first strip it (remove spaces from left and right) and then use replace method to replace the string with ""
output14 = input14.strip().replace(".", "").replace(",", "").replace(" ", "")

# YOUR CODE HERE
print(output14)
```

Output:
```
incomprehensibilities
```

<!-- grade_id: 14-1_String_character_replacement, points: 5, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 15 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(5 Points) 
Mentioned below is an extract taken from the following Wikipedia page: https://en.wikipedia.org/wiki/Python_(programming_language)

> Python is an interpreted high-level general-purpose programming language. Its design philosophy emphasises code readability with its use of significant indentation. Its language constructs as well as its object-oriented approach aim to help programmers write clear, logical code for small and large-scale projects. Python is dynamically-typed and garbage-collected. It supports multiple programming paradigms, including structured (particularly, procedural), object-oriented and functional programming. It is often described as a "batteries included" language due to its comprehensive standard library. Python consistently ranks as one of the most popular programming languages.

First create a variable `language_details` and assign the the above-mentioned extract to it.


<!-- grade_id: question_15_setup, points: 0, grade: false, solution: true -->
```python
# Create your variable here
#we can use triple quotes to assign big strings like above 
language_details = """Python is an interpreted high-level general-purpose programming language. Its design philosophy emphasises code readability with its use of significant indentation. Its language constructs as well as its object-oriented approach aim to help programmers write clear, logical code for small and large-scale projects. Python is dynamically-typed and garbage-collected. It supports multiple programming paradigms, including structured (particularly, procedural), object-oriented and functional programming. It is often described as a "batteries included" language due to its comprehensive standard library. Python consistently ranks as one of the most popular programming languages.
"""
print(language_details)
```

Output:
```
Python is an interpreted high-level general-purpose programming language. Its design philosophy emphasises code readability with its use of significant indentation. Its language constructs as well as its object-oriented approach aim to help programmers write clear, logical code for small and large-scale projects. Python is dynamically-typed and garbage-collected. It supports multiple programming paradigms, including structured (particularly, procedural), object-oriented and functional programming. It is often described as a "batteries included" language due to its comprehensive standard library. Python consistently ranks as one of the most popular programming languages.
```

Now write a program to do the following:

1. Check if the phrase `dynamically-typed` is in `language_details`. Create a variable `output15_i` and assign the result to it. [Expected output: True]
2. Check if the phrase `object-oriented` is in `language_details`. Create a variable `output15_ii` and assign the result to it. [Expected output: True]
3. Check if the word `Java` is **NOT** in `language_details`. Create a variable `output15_iii` and assign the result to it. [Expected output: True] 

Try your program with different input text, such as this extract about Java (from https://en.wikipedia.org/wiki/Java_(programming_language) )

> Java is a high-level, class-based, object-oriented programming language that is designed to have as few implementation dependencies as possible.

***

<details>
   <summary><span style="color:blue">Click to see Hints</span></summary>

    * To check if a given phrase or character is present in a string, use the keyword `in`.
    * To check if a given phrase or character is NOT present in a string, use the keyword `not in`.
</details>

<!-- grade_id: question_15, points: 0, grade: false, solution: true -->
```python
# Write your solution here
output15_i = 'dynamically-typed' in language_details
output15_ii = 'object-oriented' in language_details
output15_iii = "Java" not in language_details
#Checking dynamically-typed in the string or not 
print(f"Checking dynamically-typed in language_details or not: {output15_i}") #dynamically-typed in language_details
print(f"Checking object-oriented in language_details or not: {output15_ii}") #object-oriented in language_details
print(f"Checking Java not in language_details or not: {output15_iii}") #Java is not in language_details
```

Output:
```
Checking dynamically-typed in language_details or not: True
Checking object-oriented in language_details or not: True
Checking Java not in language_details or not: True
```

<!-- grade_id: 15-1_String_checks, points: 5, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 16 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(5 Points)
Using string methods, swap the `, ` and `.` in a given string. Create a variable `output16` and assign the solution to it.

For example for the input `10, 612, 419.001` the output would become `10.612.419, 001`. 

```python
input16 = "10, 612, 419.001"
```

<!-- grade_id: question_16, points: 0, grade: false, solution: true -->
```python
# Write your solution here
#to do this - we will first replace the `,` with a `#` so that we can use it for later 
#then we will replace all the `.` with `,` 
#lastly we will replace `#` with back to `.`

output16 = input16.replace(', ', '#').replace('.', ', ').replace('#', '.')
# YOUR CODE HERE
print(f"Input: {input16}\nOutput: {output16}")
```

Output:
```
Input: 10, 612, 419.001
Output: 10.612.419, 001
```

<!-- grade_id: 16-1_String_swap, points: 5, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 17 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(6 Points)
Given the sentence `hEY! i DON'T REMEMBER YOUR nAME, sORRY! wHAT WAS IT AGAIn?`: 

1. Make the lower case letters upper case and the upper case letters lower case in the above sentence. Create a variable `correct_txt` and assign the solution to it. [Expected Output: Hey! I don't remember your Name, Sorry! What was it agaiN?]
<details>
  <summary><span style="color:blue">Click to see Hints</span></summary>
  <ul><li>
    Review the <a href="https://docs.python.org/3/library/stdtypes.html#string-methods" >python built-in types documentation for String methods</a> for a method that swaps the case of all letters in a string...</li>
    </ul>
</details>

2. Find the number of times the letter `n` appears in the `correct_txt` using string methods and assign it to the variable `count_n`
3. Change `correct_txt` to lowercase. Now the string in `correct_txt` should be all lowercase. Count the number of times the letter `n` appears in it. Create variable `count_n_new` and assign the solution to it.
4. Set the variable `user_name` to David. Append the `user_name` three times (using string operations) **on a new line** of `correct_txt` and assign it to the `updated_txt` variable.  
  The `updated_txt` should looks like:  
```
hey! i don't remember your name, sorry! what was it again?
David David David
```

5. Convert all the letters in `updated_txt` to uppercase, and assign the solution to `output17_i` variable.
6. Check whether the word "SORRY" is in `output17_i` using string methods, and assign the solution to `output17_ii` variable. [Expected output: True]
<details>
  <summary><span style="color:blue">Click to see Hints</span></summary>
  <ul><li>
      You will need to use the keyword `in` to perform this, as you saw in question 15.</li>
    </ul>
</details>

7. Make sure everything still works as expected if you change the input sentence.

```python
input17 = "hEY! i DON'T REMEMBER YOUR nAME, sORRY! wHAT WAS IT AGAIn?"
```

<!-- grade_id: question_17, points: 0, grade: false, solution: true -->
```python
# Write your solution here
#1.
correct_txt = input17.swapcase() #swapcase will convert lower-case to upper-case and vica versa 
#2.
count_n = correct_txt.count('n') #using count method to calculate the occurence of a specific element
#3.
correct_txt = correct_txt.lower()
count_n_new = correct_txt.count('n') #using lower method to convert to lowercase
#4.
user_name = "David"
updated_txt = correct_txt + "\n" +  user_name + " " + user_name + " " + user_name
#5.
output17_i = updated_txt.upper()
#6
output17_ii = "SORRY" in output17_i
# YOUR CODE HERE
print(f"Solution17-a: {correct_txt}")
print(f"Solution17-b: {count_n}")
print(f"Solution17-c: {count_n_new}")
print(f"Solution17-d: {updated_txt}")
print(f"Solution17-e: {output17_i}")
print(f"Solution17-f: {output17_ii}")
```

Output:
```
Solution17-a: hey! i don't remember your name, sorry! what was it again?
Solution17-b: 1
Solution17-c: 3
Solution17-d: hey! i don't remember your name, sorry! what was it again?
David David David
Solution17-e: HEY! I DON'T REMEMBER YOUR NAME, SORRY! WHAT WAS IT AGAIN?
DAVID DAVID DAVID
Solution17-f: True
```

<!-- grade_id: 17-1_String_cases, points: 6, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 18 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(6 Points)
In addition to the *int* and *float* that we learnt in Topic 1, Python also has one other numeric type called
*complex numbers*. Complex numbers are written with a "j" as the imaginary part (e.g., *3+5j*).
1. Create a variable `myfirstcomplex` and assign the value `2+5j` to it.
2. Given an int `input18_int`, convert the int to a complex number and assign the solution to the `output18_i` variable. Print the type of `output18_i` to verify that it's a complex number.
3. Given an int `input18_float`, convert the float to a complex number and assign the solution to the `output18_ii` variable. Print the type of `output18_ii` to verify that it's a complex number.

***
<details>
  <summary><span style="color:blue">Click to see Hints</span></summary>
    <ul>
        <li>Python supports three types of numbers – int, float, and complex.</li>
        <li>`complex()` to convert an integer or a float into a complex number in Python.</li></ul>

</details>


```python
input18_int = 5
input18_float = 17.81
```

<!-- grade_id: question_18, points: 0, grade: false, solution: true -->
```python
# Write your solution here
myfirstcomplex = complex(2, 5)
output18_i = complex(input18_int)
output18_ii = complex(input18_float)
# YOUR CODE HERE
print(f"my first complex number: {myfirstcomplex}")
print(f"conversion of {input18_int} to complex number: {output18_i}: Datatype: {type(output18_i)}")
print(f"conversion of {input18_float} to complex number: {output18_ii}: Datatype: {type(output18_ii)}")
```

Output:
```
my first complex number: (2+5j)
conversion of 5 to complex number: (5+0j): Datatype: <class 'complex'>
conversion of 17.81 to complex number: (17.81+0j): Datatype: <class 'complex'>
```

<!-- grade_id: 18-1_Complex_numbers, points: 6, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 19 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(6 Points)
Write a Python program to convert some number of hours from 24 hour format to 12 hour format as described below. 
1. What is `20` o'clock of 24 hours clock in 12 hours clock. Create a variable `output19_i` and assign the solution to it. [Expected output: 8]
2. Update your program to handle negative hours by removing the negative mark programmatically since time cannot be negative.  
  Test that it works by changing `hours` to `-19`.
3. Check if the time from your conversion in `output19_i` is an odd number.  
  Make sure the output is a boolean type, using the bool() function to convert the result.
  Create a variable `output19_ii` and assign the boolean value to it. [Expected output: True]
***
<details>
  <summary><span style="color:blue">Click to see Hints</span></summary>
    
1. You will need to use the modulo operation (`%`) to convert the time from 24 hours clock to 12 hours clock as well as to check if the time is odd. The modulo operation is an arithmetic operation that finds the remainder of the division of one number by another. The remainder is called the modulus of the operation. <br>
e.g., 9 divided by 2 equals 4, with a remainder of 1
    
2. To remove the negative mark, you will need to use built-in `abs()` function.
</details>

```python
hours = 20
```

<!-- grade_id: question_19, points: 0, grade: false, solution: true -->
```python
# Write your solution here
hours = abs(hours)
# YOUR CODE HERE
converted_hours = hours % 12 
if not converted_hours:
    converted_hours = 12 
output19_i = converted_hours
print(output19_i)

hours = -19 
print(abs(hours) == 19)

output19_ii = bool(output19_i % 2 != 0) #checking converted hours is odd or not (True for odd)
print(output19_ii)
```

Output:
```
8
True
False
```

<!-- grade_id: 19-1_24_Hour_time, points: 6, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 20 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(6 Points)
Write a Python progarmme to convert time seconds, minutes, hours and days as described below. 

1. Convert a given number of seconds to minutes and seconds (e.g., `191 seconds` is `3` minutes and `11` seconds). Create variables `output20_i_minutes` and `output_i_seconds` and assign the number of minutes (int) and number of seconds (int) to these two variables, respectively.

2. Convert a given number of hours to days and hours. Create variables `output20_ii_days` and `output20_ii_hours` and assign the number of days (int) and number of hours (int) to these variables, respectively.

<details>
  <summary><span style="color:blue">Click to see Hints</span></summary>
    
1. You will need to use the integer division operation to obtain the number of minutes. 

2. You will need to use the modulo operation to obtain the number of seconds.
</details>

```python
seconds = 303
hours = 1909
```

<!-- grade_id: question_20, points: 0, grade: false, solution: true -->
```python
# Write your solution here
output20_i_minutes = seconds // 60 
output20_i_seconds = seconds % 60 
print(f"Output for minutes: {output20_i_minutes} and seconds: {output20_i_seconds}")

output20_ii_days = hours // 24 
output20_ii_hours = hours % 24 


# YOUR CODE HERE
print(f"Output for days: {output20_ii_days} and hours: {output20_ii_hours}")
```

Output:
```
Output for minutes: 5 and seconds: 3
Output for days: 79 and hours: 13
```

<!-- grade_id: 20-1_Seconds_conversion, points: 6, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 21 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(6 Points)
Given the following string/sentence: 
```
'          \n   \t   In this week we learnt about "Environments" and "Variables & Types" in Python! \n \n     '
```

1. Remove all whitespace at the beginning and at the end of the string using built-in functions, and assign the solution to `output21_i`.  
  Check your work using the string length of `input21` and `output21_i` (*length of `input21` is 105 and length of `output21_i` should be 78*).
2. In `input21`, remove all whitespace **only** at the beginning of the string using built-in functions, and assign the solution to `output21_ii`. Again, check your work using the string lengths (*length of `output21_ii` should be 87*).
3. In `input21`, remove all whitespace **only** at the end of the string using built-in functions, and assign the solution to `output21_iii`. Again, check your work using the string lengths (*length of `output21_iii` should be 96*).
4. In `input21`, replace all whitespace with `&` using built-in functions, and assign the solution to `output21_iv`. Your output should look like: `&&&&&&&&&&&&&&&&&&In&this&week&we&learnt&about&"Environments"&and&"Variables&&&Types"&in&Python!&&&&&&&&&`

<details>
  <summary><span style="color:blue">Click to see Hints for 1, 2, 3 & 4</span></summary>

* `strip()` function removes any leading (spaces at the beginning) and trailing (spaces at the end) characters (space is the default leading character to remove).
* `lstrip()` function removes any leading characters (space is the default leading character to remove).
* `rstrip()` function removes any trailing characters (space is the default leading character to remove).
* use `replace()` to perform this.
    
If you need more information on them, use `?` or `help()`.

</details>

5. Check if `input21` starts with `In`. Create the variable `output21_v` and assign the solution to it.
6. Check if `output21_i` ends with `Python!`. Create the variable `output21_vi` and assign the solution to it.

<details>
  <summary><span style="color:blue">Click to see Hints for 5 & 6</span></summary>

* You will need to use startswith() and endswith() methods. If you need more information on them, use `?` or `help()`.

</details>

```python
input21 = '          \n   \t   In this week we learnt about "Environments" and "Variables & Types" in Python! \n \n     '
```

<!-- grade_id: question_21, points: 0, grade: false, solution: true -->
```python
# Write your solution here
#1. remove all white-spaces using strip method() (this will remove from left and right both )
output21_i = input21.strip() 
#2. remove all white-spaces from left side of the string we will use lstrip() 
output21_ii = input21.lstrip()
#3. remove all white-spaces from right side of the string we will use rstrip() 
output21_iii = input21.rstrip() 
#4. to replace all whitespaces with & we will use .replace() method 
output21_iv = input21.replace(" ", "&").replace("\t", "&").replace("\n", "&")
#5. check if input21 have In in starting we can use `startswith()` method
output21_v = input21.startswith("In")
#6. check if input21 have Python! in the end we can use 
output21_vi = output21_i.endswith("Python!")
# YOUR CODE HERE
print(len(input21), len(output21_i)) 
print(len(input21), len(output21_ii)) 
print(len(input21), len(output21_iii)) 
print(output21_iv)
print(output21_v)
print(output21_vi)
```

Output:
```
105 78
105 87
105 96
&&&&&&&&&&&&&&&&&&In&this&week&we&learnt&about&"Environments"&and&"Variables&&&Types"&in&Python!&&&&&&&&&
False
True
```

<!-- grade_id: 21-1_String_whitespace, points: 6, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 22 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(6 Points)
Write a Python Program to calculate the following.
1. Write a python program to generate three random int numbers between 1 and 10000 and assign them to the variables `output22_num1`, `output22_num2` and `output22_num3`.
2. Use these three numbers to sort them in increasing order. To do so, first find the minimum and maximum numbers using min() and max() methods. Create the two variables `output22_min` and `output22_max` and assign the solutions to them. Then find the middle number using arithmetic operations and assign the solution to `output22_mid`. From the above variables derived the following string: `The numbers in increasing order: min_num, mid_num, max_num` and assign it to the variable `output22`. 


*(Note: `min_num`, `midd_num`, `max_num` in the given string are the numbers you derived in `output22_min`, `output22_mid`, `output22_max`)*.

***
<details>
  <summary><span style="color:blue">Click to see Hints</span></summary>
    
* To find the middle number using arithmetic operations, you will need to get the sum of the three numbers and reduce the minimum and maximum nubers from it. e.g., sum_of_three_numbers - min_num - max_num
</details>

<!-- grade_id: question_22, points: 0, grade: false, solution: true -->
```python

# Write your solution here
import random
output22_num1, output22_num2, output22_num3 = random.randint(1, 10000), random.randint(1, 10000), random.randint(1, 10000)


output22_min, output22_max  = min(output22_num1, output22_num2, output22_num3), max(output22_num1, output22_num2, output22_num3)

sum_of_three_numbers = output22_num1 + output22_num2 + output22_num3
output22_mid = sum_of_three_numbers - output22_min - output22_max


# YOUR CODE HERE
output22 = f"The numbers in increasing order: {output22_min}, {output22_mid}, {output22_max}"
print(output22)
```

Output:
```
The numbers in increasing order: 570, 5737, 7740
```

<!-- grade_id: 22-1_Random_numbers, points: 6, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```

#### Question 23 &nbsp;&nbsp;&nbsp;&nbsp;&nbsp;&nbsp;(10 Points)
Write a Python program to convert a given number of seconds into days, hours, minutes and seconds.  
Create variables `output23_days`, `output23_hours`, `output23_minutes` and `output23_seconds` and assign the number of days (int), number of hours (int), number of minutes (int) and number of seconds (int) to these variables, respectively. 

```python
seconds = 219991
```

<!-- grade_id: question_23, points: 0, grade: false, solution: true -->
```python
# Write your solution here
output23_days = seconds // (24 * 3600)
remaining_seconds = seconds % (24 * 3600)
output23_hours = remaining_seconds // 3600
remaining_seconds %= 3600

output23_minutes = remaining_seconds // 60
output23_seconds = remaining_seconds % 60




# YOUR CODE HERE
print(f"Days: {output23_days}, Hours: {output23_hours}, Minutes: {output23_minutes}, Seconds: {output23_seconds}")
```

Output:
```
Days: 2, Hours: 13, Minutes: 6, Seconds: 31
```

<!-- grade_id: 23-1_Seconds_coversion_2, points: 10, grade: true, solution: false -->
```python
# Testing Cell (Do NOT not modify this cell)
```